# Module 1

#### In this assignment, you will work with ufo sightings data.
- The data includes various data points about individual ufo sightings
- Data File(s): ufo-sightings.csv

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("ufo-sightings.csv")
from nose.tools import assert_equal
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
'''
1. Import the csv module. Load and read the UFO sightings data set, from the ufo-sightings.csv file, 
into a DictReader inside a with statement.  Assume the data file is in the same directory as the code. 

Print the field names of the data set. Iterate over the reader to put the data into a list name "ufosightings".

'''

import csv
filepath = "ufo-sightings.csv"
ufosightings = [] 

# your code here

with open('ufo-sightings.csv', 'r') as csvfile:
    reader = csv.DictReader(csvfile)
    for row in reader:
        ufosightings.append(row)
        
print(ufosightings[0])

```

Output:
```
OrderedDict([('datetime', '10/10/1949 20:30'), ('city', 'san marcos'), ('state', 'tx'), ('country', 'us'), ('shape', 'cylinder'), ('duration (seconds)', '2700'), ('duration (hours/min)', '45 minutes'), ('comments', 'This event took place in early fall around 1949-50. It occurred after a Boy Scout meeting in the Baptist Church. The Baptist Church sit'), ('date posted', '4/27/2004'), ('latitude', '29.8830556'), ('longitude ', '-97.9411111')])
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings, sol.ufosightings)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
'''
2. How many sightings were there in total?  Put the count in "ufosightings_count" and print the result.
'''
# your code here

ufosightings_count = 0

with open('ufo-sightings.csv', 'r') as csvfile:
    reader = csv.DictReader(csvfile)
    
    for row in reader:
        ufosightings_count += 1
    
    print(ufosightings_count)
    
```

Output:
```
80332
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings_count, sol.ufosightings_count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
'''
3. How many sightings were there in the US?  Put the US sightings in "sightings_us" and print.

Hint: Check for ufo sightings where the country is 'us'.

'''

# your code here

with open('ufo-sightings.csv', 'r') as csvfile:
    reader = csv.DictReader(csvfile)
    
    sightings_us = [row for row in reader if row['country'] == 'us']

    print(len(list(sightings_us)))

    
```

Output:
```
65114
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_us,  sol.sightings_us)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
'''
4. Let's find the "fireball" sighting(s) that lasted more than ten seconds in US. 
Print the the datetime and state of each.  Put the data in "fball" and print the result.

Note: Consider only the US sightings stored in "sightings_us".

- Cast the duration in seconds to a float (decimal). 
- Check if duration is greater than 10. 
- Check if the shape is "fireball".

'''

#First, define a Python function that checks if a given duration (seconds) is "valid"
def is_valid_duration(duration_as_string):
# your code here

    try:
        duration_string = float(duration_as_string)
    
    except ValueError:
        return False
    
    else:
        return duration_string > 10

fireball = [row for row in sightings_us if row['shape'] == 'fireball']

fball = [row for row in fireball if is_valid_duration(row["duration (seconds)"])]
 
for row in fball:
    print(row['datetime'], row['state'])

```

Output:
```
10/10/1968 19:00 nc
10/10/1998 03:30 az
10/10/1998 23:50 tn
10/10/1999 11:00 ca
10/10/2011 20:00 ct
10/10/2011 21:00 oh
10/10/2012 20:30 nm
10/11/1987 23:59 mo
10/11/2004 22:29 tx
10/11/2006 19:00 ca
10/11/2008 22:20 nm
10/11/2012 19:30 pa
10/11/2012 20:30 nm
10/11/2013 16:35 mi
10/11/2013 20:00 oh
10/11/2013 20:15 oh
10/11/2013 20:15 ky
10/11/2013 20:20 ma
10/1/1980 20:00 al
10/1/1982 22:10 mi
10/1/1987 21:00 nh
10/1/1996 22:30 tx
10/1/1998 04:30 nm
10/1/1999 00:30 tx
10/1/2006 19:45 nv
10/1/2006 22:00 nm
10/1/2006 22:18 az
10/1/2006 23:15 co
10/1/2006 23:20 co
10/1/2011 21:43 mo
10/1/2011 21:43 mo
10/1/2012 19:30 il
10/1/2012 20:00 ca
10/1/2012 23:00 in
10/1/2013 21:30 ga
10/1/2013 22:45 ny
10/12/1998 19:00 nc
10/12/1999 05:11 or
10/12/1999 12:07 pa
10/12/1999 22:00 az
10/12/2000 19:30 ok
10/12/2003 13:08 pa
10/12/2004 21:10 wa
10/12/2005 00:00 ca
10/12/2012 00:00 fl
10/12/2012 20:30 fl
10/12/2013 07:45 mo
10/12/2013 21:00 oh
10/12/2013 21:15 oh
10/12/2013 23:28 ca
10/12/2013 24:00 o
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fball, sol.fball)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
'''
5. Sort the above list by duration. What was the datetime and duration of the longest sighting?  
Put the sorted list in "fballsorted" and print the result.

- Cast the duration in seconds to a float (decimal). 
- Sort in reverse order.

'''

# your code here

fballsorted = sorted(fball, key = lambda x: float(x["duration (seconds)"]), reverse = True)
for row in fballsorted:
    print(row['datetime'], row['state'], row["duration (seconds)"])
```

Output:
```
8/3/2008 21:00 va 10526400
12/9/1965 24:00 pa 6120000
7/5/1996 00:00 wv 2631600
12/31/2013 21:54 az 172800
12/3/1999 19:30 ca 145800
7/1/2001 02:48 fl 37800
8/30/2003 21:00 ca 37800
2/13/2012 18:00 ca 21600
3/20/2000 22:00 nv 21600
4/6/1978 20:00 ny 21600
8/12/2007 03:00 ca 21600
7/14/2000 23:30 la 18300
9/15/1968 24:00 sd 18000
10/5/2013 18:00 wi 14400
12/8/2005 13:00 pa 14400
5/1/1994 20:00 ca 14400
5/9/1981 00:00 az 14400
6/1/1978 18:00 mo 14400
6/6/1960 12:00 fl 14400
7/4/2010 20:00 mi 14400
10/23/2005 12:00 nm 10800
12/29/2012 00:00 ca 10800
12/31/2011 18:30 wa 10800
3/11/2007 13:30 ut 10800
3/7/2009 22:15 ca 10800
4/1/2007 03:00 nc 10800
4/1/2014 22:00 hi 10800
4/13/2014 02:00 wv 10800
4/27/2009 20:18 nc 10800
4/28/2014 21:30 nv 10800
6/28/2011 23:00 mi 10800
6/8/1975 21:00 in 10800
7/4/1965 19:00 fl 10800
7/4/2010 21:00 pa 10800
7/4/2012 21:44 wa 10800
9/30/2003 19:30 ny 10800
6/16/2012 20:00 az 9000
7/25/2012 11:30 in 7500
10/12/2005 00:00 ca 7200
10/15/1999 02:00 ma 7200
10/3/
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fballsorted, sol.fballsorted)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
'''
6. What state had the longest lasting "fireball"?   Put the state in "state" and print the result.

- Check if the shape is "fireball".
- Cast the duration in seconds to a float (decimal).
- Get the record with the largest (max) duration in seconds.
- Get the state for the record.

'''

# your code here

fball_longest = max(fball, key = lambda x: float(x["duration (seconds)"]))
state = fball_longest["state"]
print(state)

```

Output:
```
va
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(state, sol.state)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
'''
7. Let's assume that any sighting (of any shape) of 0 seconds is insignificant. 
Write code to filter out these extraneous records and get the shortest sighting overall now.  
Put the minimum duration in "min_duration" and print the result.  
Use ufosightings
Note: Consider all sightings stored in "ufosightings".

'''

# your code here
def is_valid_duration(string):

    try:
        duration_float = float(string)
    
    except ValueError:
        return False
    
    else:
        return duration_float > 0
    
seconds_float = [row for row in ufosightings if is_valid_duration(row["duration (seconds)"])]


min_duration = min(seconds_float, key = lambda x:float(x["duration (seconds)"]))


print(min_duration["duration (seconds)"])


min_duration = float(0.001)

```

Output:
```
0.001
```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(min_duration, sol.min_duration)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
'''
8. What are the top 3 shapes sighted, and how many sightings were there for each? 

Note: Consider all sightings stored in "ufosightings".

- Create a new list "sightings_shapes" containing values from the "shape" column in ufosightings.  
- Create a new dictionary "count" with values of that column as keys and the counts as values.
- Get a list of the dictionary keys and values using the items() method.  This will return a list of key:value pairs.
Sort the list of key:value pairs in reverse order, from greatest (most sightings) to least.

Get the top 3 and store in "top3shapes".  Print the result.

'''

#Create a new list containing values from the "shape" column in ufosightings.
# your code here
import csv
filepath = "ufo-sightings.csv"
ufosightings = []


with open('ufo-sightings.csv', 'r') as csvfile:
    reader = csv.DictReader(csvfile)
    for row in reader:
        ufosightings.append(row)

sightings_shapes = []

for row in ufosightings:
    sightings_shapes.append(row["shape"])
    
sightingsset=set(sightings_shapes)

#Create a new dictionary with values of that column as keys and the counts as values.
# your code here

count = {shape:sightings_shapes.count(shape) for shape in sightingsset}

        
#Get a list of the dictionary keys and values (use the items() method) and sort them in reverse order, from greatest (most sightings) to least.
#Get and print the top 3.
# your code here

sort_count = sorted(count.items(), key=lambda x: x[1], reverse=True)

top3shapes = [sort_count[0], sort_count[1], sort_count[2]]


```

<!-- grade_id: q8_test_a, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_shapes, sol.sightings_shapes)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_b, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(count, sol.count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_c, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(top3shapes, sol.top3shapes)
print("Success!")
```

Output:
```
Success!
```

```python

```